Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

NAME = ""
COLLABORATORS = ""

---

```scala
import $file.hw6stdlib
import hw6stdlib._
```

# Homework 6

In this assignment we will write the first full interpreter for Lettuce. Remember that lettuce is a functional language with let bindings. It is very similiar to a language called ML. Here we want to implement the interpreter for the language from the ground up. We will use most of the existing standard library we have developed while we write this, especially the equality functions and Dictionary data type from the last homework.

## Part 1 - The Values

Recall that we define interpreters as functions that take in expressions in the form of abstract syntax and give a value as output. In symbols:

$$
eval : Expr \rightarrow Value
$$

It follows that we will need to create a datatype to represent the values that may be computed by lettuce. Below is the grammar that defines lettuce values. Implement this as a `sealed trait` in Scala below:

$$
\begin{align}
\textbf{Value} ::=&\ NumVal\ \mathbb{Z}\\
\mid&\ BinVal\ \mathbb{B}\\
\mid&\ Error
\end{align}
$$

<!-- grade_id: values-adt, points: 0, grade: false, solution: true -->
```scala
??? // YOUR CODE HERE
```

If your definition was correct then the values in the cell below should compile.

<!-- grade_id: values-adt-test, points: 5, grade: true, solution: false -->
```scala
val v1 = NumVal(Positive(Succ(Succ(Zero))))
val v2 = BinVal(False)
val v3 = Error
```

## Part 2 - The Expressions

There are many possible expressions in the Lettuce language. We also use an abstract syntax to represent this. Here is the grammar to remind you of all the syntactic elements in Lettuce:

$$\begin{array}{rcll}
\mathbf{Expr} & ::= & Const(\mathbb{Z}) \\
 & | & Bin(\mathbb{B}) \\
 & | & Ident(\mathbf{String}) \\
 & | & Plus(\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Minus(\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Mult (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Pow (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & Neg (\mathbf{Expr}) \\
 & | & Eq (\mathbf{Expr}, \mathbf{Expr}) \\
 & | & And ( \mathbf{Expr}, \mathbf{Expr} ) \\
 & | & Or ( \mathbf{Expr}, \mathbf{Expr} ) \\
 & | & IfThenElse(\mathbf{Expr}, \mathbf{Expr}, \mathbf{Expr}) \\
 & | & Let( \mathbf{String}, \mathbf{Expr}, \mathbf{Expr}) \\
\end{array}$$

Now define a `sealed trait` for expressions in Lettuce:

<!-- grade_id: expr, points: 0, grade: false, solution: true -->
```scala
??? // YOUR CODE HERE
```

If you defined `Expr` properly then the following expressions should compile:

<!-- grade_id: expr-test, points: 5, grade: true, solution: false -->
```scala
val ex1 = Const(Positive(five))
val ex2 = Const(Negative(three))
val ex3 = Bin(True)
val ex4 = Ident("x")
val ex5 = Plus(ex1, ex2)
val ex6 = Let("x", ex5, Mult(ex4, ex4))
```

## Part 3 - The Interpreter 

We now need to define the interpreter for this language. It is a function of the form:

$$
eval : Expr \rightarrow Value
$$

Except, now that we have an execution environment($\sigma$), we need some way to include $\sigma$ in this function. This will be as an additional argument:

$$
eval : \sigma \rightarrow Expr \rightarrow Value
$$

We will let $\sigma$ be a `Dictionary` as we defined in last week's homework. To add something to a dicitonary use the `Entry` constructor to add on a new element. This should be very similiar to the way `Cons` works on lists.

Below we have given you the skeleton of the function `eval`. Fill in each case for the interpreter so that it can interpret all possible Lettuce expressions. The inference rules we covered in the classroom should be very helpful with this.

Hint: It may be helpful to define some helper functions for the inference rules that have similar forms. These are the rules like binary operations for arithmetic, bool, etc. Just like we lumped some of these rules together into a single rule, we can do the same for our interpreter in the form of auxillary functions. The bonus `eval_bin` function from Homework 4 may be a good place for inspiration.

Hint: Be ready to write some nested case matches. Most of your cases should have two nestings. One for the Sytnax element and another for evaluating the expressions it may contain. This is not true for all expressions but most will follow this pattern to some extent. Note that `Ident` has a different kind of case match because it needs to interact with the environment(a Dictionary)

<!-- grade_id: interpreter, points: 0, grade: false, solution: true -->
```scala
def eval(env : Dictionary[String, Value], expr : Expr) : Value =
??? // YOUR CODE HERE
```

If this was defined correctly then the definitions above should evaluate:

<!-- grade_id: interpreter-tests, points: 20, grade: true, solution: false -->
```scala
//val ex1 = Const(Positive(five))
//val ex2 = Const(Negative(three))
//val ex3 = Bin(True)
//val ex4 = Ident("x")
//val ex5 = Plus(ex1, ex2)
//val ex6 = Let("x", ex5, Mult(ex4, ex4))


assert(eval(EmptyDict, ex3) equals BinVal(True))
assert(eval(EmptyDict, ex4) equals Error)
assert(eval(EmptyDict, ex5) equals NumVal(Positive(two)))
assert(eval(EmptyDict, ex6) equals NumVal(Positive(four)))
```