Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```scala
NAME = ""
COLLABORATORS = ""
```

---

```scala
import $file.hw3stdlib
import hw3stdlib._
```

# Homework 3

Due Tuesday, June 18th at 11:59pm

## Problem 1 (5 pts)

Write the filter function. Remember that this should be polymorphic. So:

$$
filter : (a \rightarrow \mathbb{B}) \rightarrow List\ a \rightarrow List\ a
$$

If the predicate(The function we give it) is true for an element, then that element will be in the output list.

<!-- grade_id: filter, points: 0, grade: false, solution: true -->
```scala
def filter[A](p : (A => Bool), xs : List[A]) : List[A] =
??? // YOUR CODE HERE
```

<!-- grade_id: filter-tests, points: 5, grade: true, solution: false -->
```scala
//Tests to validate your answer
val lt2 : Int => Bool = {
    case x if x < 2 => True
    case _ => False
}
assert(filter(lt2, Empty) == Empty)
assert(filter(lt2, Cons(3,Empty)) == Empty)
assert(filter(lt2, Cons(1,Empty)) == Cons(1,Empty))
assert(filter(lt2, Cons(1,Cons(3,Empty))) == Cons(1,Empty))
assert(filter(lt2, Cons(5,Cons(0, Cons(1,Cons(3,Empty))))) == Cons(0,Cons(1,Empty)))
```

## Problem 2 (5 pts)

Write a polymorphic version of maybe. Here is a reminder of the definition we gave in class:

$$
\begin{align}
\text{Maybe}\ a ::=&\ \text{Nothing}\\
\mid&\ \text{Just}\ a
\end{align}
$$ 

Hint: you should use the type `scala.Nothing` in your definition of `Nothing`. `scala.Nothing` is a built-in "bottom" type in Scala, meaning that it is a subtype of all other types.  Don't be confused by the name collision though -- `scala.Nothing` is very different from the `Nothing` type we're asking you to define.

<!-- grade_id: maybe, points: 0, grade: false, solution: true -->
```scala
??? // YOUR CODE HERE
```

<!-- grade_id: maybe-tests, points: 5, grade: true, solution: false -->
```scala
//Tests to validate your answer
val x : Maybe[Bool] = if((new scala.util.Random()).nextBoolean()) Just(True) else Nothing
```

## Problem 3 (5 pts)

Write a higher-order function that applies a function to the element inside of $Just$ or if it is $Nothing$ maps to $Nothing$. We will call this $\texttt{map_maybe}$.

$$
\text{map_maybe} : (a \rightarrow b) \rightarrow \text{Maybe}\ a \rightarrow \text{Maybe}\ b
$$

<!-- grade_id: fmap, points: 0, grade: false, solution: true -->
```scala
def map_maybe[A, B](f : (A => B), m : Maybe[A]) : Maybe[B] = 
??? // YOUR CODE HERE
```

<!-- grade_id: fmap-tests, points: 5, grade: true, solution: false -->
```scala
//Tests to validate your answer
assert(map_maybe(lt2,Nothing) == Nothing)
assert(map_maybe(lt2,Just(0)) == Just(True))
assert(map_maybe({x:Bool => "hello world"},Just(True)) == Just("hello world"))
```