Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```scala
NAME = ""
COLLABORATORS = ""
```

---

```scala
import $file.hw2stdlib
import hw2stdlib._
```

# Homework 2

Due 6/14 at 11:59pm



# Problem 1 (5 points)

Write the minus function, which subtracts a natural number from another natural number. If the result would be less than zero then leave the answer as zero.

<!-- grade_id: minus, points: 0, grade: false, solution: true -->
```scala
def minus(x : Nat, y : Nat) : Nat = 
??? // YOUR CODE HERE
```

<!-- grade_id: minus-tests, points: 5, grade: true, solution: false -->
```scala
assert(minus(three, two) == one)
assert(minus(five,one) == four)
assert(minus(one,Zero) == one)
assert(minus(two,Zero) == two)
```

# Problem 2 (5 points)

Define the $\texttt{pow}$ function, which raises a natural number to some power. Notationally: 

$$
f\ x\ y = x^y
$$

Hint: Use the mult function from the stdlib for help

<!-- grade_id: pow, points: 0, grade: false, solution: true -->
```scala
def pow(x : Nat, y : Nat) : Nat =
??? // YOUR CODE HERE
```

<!-- grade_id: pow-tests, points: 5, grade: true, solution: false -->
```scala
assert(pow(one,one) == one)
assert(pow(one,five) == one)
assert(pow(three,two) == nine)
assert(pow(two,three) == eight)
assert(pow(ten,Zero) == one)
```

# Problem 3 (5 points)

Write the "less than or equal to" function for numbers. This function should return $True$ if $x < y$ and $False$ otherwise.

<!-- grade_id: lte, points: 0, grade: false, solution: true -->
```scala
def lte(x : Nat, y : Nat) : Bool =
??? // YOUR CODE HERE
```

<!-- grade_id: lte-tests, points: 5, grade: true, solution: false -->
```scala
assert(lte(one,one) == True)
assert(lte(four,four) == True)
assert(lte(two,five) == True)
assert(lte(five,two) == False)
```